---
title       : How to maximize likelihood?
description : We want to take it to the max!

--- type:NormalExercise lang:r xp:100 skills:1 key:f4631fbfdb
## Loading up Stan


*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```
